*version 190901.0*

# Assignment 2
For this assignment you'll be looking at 2017 data on immunizations from the CDC. Your datafile for this assignment is in [assets/NISPUF17.csv](assets/NISPUF17.csv). A data users guide for this, which you'll need to map the variables in the data to the questions being asked, is available at [assets/NIS-PUF17-DUG.pdf](assets/NIS-PUF17-DUG.pdf).

## Question 1
Write a function called `proportion_of_education` which returns the proportion of children in the dataset who had a mother with the education levels equal to less than high school (<12), high school (12), more than high school but not a college graduate (>12) and college degree.

*This function should return a dictionary in the form of (use the correct numbers, do not round numbers):* 
```
    {"less than high school":0.2,
    "high school":0.4,
    "more than high school but not college":0.2,
    "college":0.2}
```


<!-- grade_id: cell-eea16d020eb52ae7, points: 0, grade: false, solution: true -->
```python
import pandas as pd
import numpy as np

pd.set_option('display.max_columns', None)

df = pd.read_csv('assets/NISPUF17.csv')
#####################
df_EDUC1_Less_Than_12 = df.where(df['EDUC1']==1)

df_EDUC1_Equal_To_12 = df.where(df['EDUC1']==2)

df_EDUC1_More12_Not_Col_Grad = df.where(df['EDUC1']==3)

df_EDUC1_Col_Grad = df.where(df['EDUC1']==4)

######################

df_EDUC1_Series = df['EDUC1']

df_EDUC1_Less_Than_12_Series = df_EDUC1_Less_Than_12['EDUC1']

df_EDUC1_Equal_To_12_Series = df_EDUC1_Equal_To_12['EDUC1']

df_EDUC1_More12_Not_Col_Grad_Series = df_EDUC1_More12_Not_Col_Grad['EDUC1']

df_EDUC1_Col_Grad_Series = df_EDUC1_Col_Grad['EDUC1']

#######################
df_EDUC1_Series = df_EDUC1_Series.dropna()
df_EDUC1_Less_Than_12_Series = df_EDUC1_Less_Than_12_Series.dropna()
df_EDUC1_Equal_To_12_Series = df_EDUC1_Equal_To_12_Series.dropna()
df_EDUC1_More12_Not_Col_Grad_Series = df_EDUC1_More12_Not_Col_Grad_Series.dropna()
df_EDUC1_Col_Grad_Series = df_EDUC1_Col_Grad_Series.dropna()

#######################

lessthanhighschoolprop = (df_EDUC1_Less_Than_12_Series.count())/(df_EDUC1_Series.count())
equaltohighschoolprop = (df_EDUC1_Equal_To_12_Series.count())/(df_EDUC1_Series.count())
morethan12notcolgradprop = (df_EDUC1_More12_Not_Col_Grad_Series.count())/(df_EDUC1_Series.count())
colgradprop = (df_EDUC1_Col_Grad_Series.count())/(df_EDUC1_Series.count())


proportion_of_education_dictionary = {"less than high school":lessthanhighschoolprop,
                              "high school":equaltohighschoolprop,
                              "more than high school but not college":morethan12notcolgradprop,
                              "college": colgradprop }
    


def proportion_of_education():
    return proportion_of_education_dictionary
    
    raise NotImplementedError()
    
    

```

<!-- grade_id: cell-c0eeef201366f51c, points: 1, grade: true, solution: false -->
```python
assert type(proportion_of_education())==type({}), "You must return a dictionary."
assert len(proportion_of_education()) == 4, "You have not returned a dictionary with four items in it."
assert "less than high school" in proportion_of_education().keys(), "You have not returned a dictionary with the correct keys."
assert "high school" in proportion_of_education().keys(), "You have not returned a dictionary with the correct keys."
assert "more than high school but not college" in proportion_of_education().keys(), "You have not returned a dictionary with the correct keys."
assert "college" in proportion_of_education().keys(), "You have not returned a dictionary with the correct keys."

```

## Question 2

Let's explore the relationship between being fed breastmilk as a child and getting a seasonal influenza vaccine from a healthcare provider. Return a tuple of the average number of influenza vaccines for those children we know received breastmilk as a child and those who know did not.

*This function should return a tuple in the form (use the correct numbers:*
```
(2.5, 0.1)
```

<!-- grade_id: cell-77f18c512324eabb, points: 0, grade: false, solution: true -->
```python
##CBF_01(1 = breastafeed 2 = no breastfeed)
##P_NUMFLU




target_dataframe = df[['CBF_01', 'P_NUMFLU']] 
##########################

breastfeed_dataframe = target_dataframe.where(target_dataframe['CBF_01']==1)

breastfeed_dataframe.dropna(inplace = True)
##########################
no_breastfeed_dataframe = target_dataframe.where(target_dataframe['CBF_01']==2)

no_breastfeed_dataframe.dropna(inplace = True)
##########################
average_PNUMFLU_breastfeed = (breastfeed_dataframe['P_NUMFLU'].sum())/(len(breastfeed_dataframe))

average_PNUMFLU_no_breastfeed =  (no_breastfeed_dataframe['P_NUMFLU'].sum())/(len(no_breastfeed_dataframe))




def average_influenza_doses():
    return (average_PNUMFLU_breastfeed, average_PNUMFLU_no_breastfeed)
    raise NotImplementedError()
```

<!-- grade_id: cell-54a3ba6cff31caa7, points: 1, grade: true, solution: false -->
```python
assert len(average_influenza_doses())==2, "Return two values in a tuple, the first for yes the second for no."

```

## Question 3
It would be interesting to see if there is any evidence of a link between vaccine effectiveness and sex of the child. Calculate the ratio of the number of children who contracted chickenpox but were vaccinated against it (at least one varicella dose) versus those who were vaccinated but did not contract chicken pox. Return results by sex. 

*This function should return a dictionary in the form of (use the correct numbers):* 
```
    {"male":0.2,
    "female":0.4}
```

<!-- grade_id: cell-a0a9e6fe67698006, points: 0, grade: false, solution: true -->
```python
###P_NUMVRC = any numberic value for number of varicella doses
###HAD_CPOX: 1=YES, 2=NO
###SEX: 1=Male 2=Female







target_dataframe = df[['HAD_CPOX', 'P_NUMVRC', 'SEX']] 
##########################
had_chickenpox = target_dataframe.where(target_dataframe['HAD_CPOX']==1)
no_chickenpox = target_dataframe.where(target_dataframe['HAD_CPOX']==2)

##########################
had_chickenpox_vaccinated = had_chickenpox.where(had_chickenpox['P_NUMVRC']>0)
no_chickenpox_vaccinated = no_chickenpox.where(no_chickenpox['P_NUMVRC'] >0)

##########################
had_chickenpox_vaccinated_male = had_chickenpox_vaccinated.where(had_chickenpox_vaccinated['SEX']==1)

no_chickenpox_vaccinated_male = no_chickenpox_vaccinated.where(no_chickenpox_vaccinated['SEX']==1)


had_chickenpox_vaccinated_female = had_chickenpox_vaccinated.where(had_chickenpox_vaccinated['SEX']==2)
no_chickenpox_vaccinated_female =  no_chickenpox_vaccinated.where(no_chickenpox_vaccinated['SEX']==2)



had_chickenpox_vaccinated_male.dropna(inplace = True)
no_chickenpox_vaccinated_male.dropna(inplace = True)
had_chickenpox_vaccinated_female.dropna(inplace = True)
no_chickenpox_vaccinated_female.dropna(inplace = True)
##########################

male_ratio = len(had_chickenpox_vaccinated_male)/len(no_chickenpox_vaccinated_male)

female_ratio = len(had_chickenpox_vaccinated_female)/len(no_chickenpox_vaccinated_female)






def chickenpox_by_sex():
    chicken_dict =    {
    "male" : male_ratio,
    "female" : female_ratio,

     }
    
    return chicken_dict

    raise NotImplementedError()
```

<!-- grade_id: cell-c4f1714db100c865, points: 1, grade: true, solution: false -->
```python
assert len(chickenpox_by_sex())==2, "Return two values in a tuple, the first for males the second for females."

```

## Question 4
A correlation is a statistical relationship between two variables. If we wanted to know if vaccines work, we might look at the correlation between the use of the vaccine and whether it results in prevention of the infection or disease [1]. In this question, you are to see if there is a correlation between having had the chicken pox and the number of chickenpox vaccine doses given (varicella).

Some notes on interpreting the answer. If the `had_chickenpox_column` is either `1` (for yes) or `2` for no, and that the `num_chickenpox_vaccine_column` is the number of doses a child has been given of the varicella vaccine, then a positive correlation (e.g. `corr > 0`) would mean that an increase in `had_chickenpox_column` (which means more no's) would mean an increase in the `num_chickenpox_vaccine_column` (which means more doses of vaccine). If `corr < 0` then there is a negative correlation, indicating that having had chickenpox is related to an increase in the number of vaccine doses. Also, `pval` refers to the probability the relationship observed is significant. In this case `pval` should be very very small (will end in `e-18` indicating a very small number), which means the result unlikely to be by chance.

[1] This isn't really the full picture, since we are not looking at when the dose was given. It's possible that children had chickenpox and then their parents went to get them the vaccine. Does this dataset have the data we would need to investigate the timing of the dose?

<!-- grade_id: cell-8afff07f564cf79a, points: 0, grade: false, solution: true -->
```python
def corr_chickenpox():
    import scipy.stats as stats
    import numpy as np
    import pandas as pd
    
    # this is just an example dataframe
    almost_pre_target_dataframe = df[['HAD_CPOX', 'P_NUMVRC']]
    pre_target_dataframe =  almost_pre_target_dataframe[(almost_pre_target_dataframe['HAD_CPOX'] >= 1) & (almost_pre_target_dataframe['HAD_CPOX'] <= 2)]
    target_dataframe = pre_target_dataframe.dropna()

    

    # here is some stub code to actually run the correlation
    corr, pval=stats.pearsonr(target_dataframe["HAD_CPOX"],target_dataframe["P_NUMVRC"])
    
    
    # just return the correlation
    return corr


    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-73408733533a29a5, points: 1, grade: true, solution: false -->
```python
assert -1<=corr_chickenpox()<=1, "You must return a float number between -1.0 and 1.0."

```